# Ordinary Differential Equations Exercise 3

## Imports

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
import seaborn as sns
from scipy.integrate import odeint
from IPython.html.widgets import interact, fixed
```

## Damped, driven nonlinear pendulum

The equations of motion for a simple [pendulum](http://en.wikipedia.org/wiki/Pendulum) of mass $m$, length $l$ are:

$$
\frac{d^2\theta}{dt^2} = \frac{-g}{\ell}\sin\theta
$$

When a damping and periodic driving force are added the resulting system has much richer and interesting dynamics:

$$
\frac{d^2\theta}{dt^2} = \frac{-g}{\ell}\sin\theta - a \omega - b \sin(\omega_0 t)
$$

In this equation:

* $a$ governs the strength of the damping.
* $b$ governs the strength of the driving force.
* $\omega_0$ is the angular frequency of the driving force.

When $a=0$ and $b=0$, the energy/mass is conserved:

$$E/m =g\ell(1-\cos(\theta)) + \frac{1}{2}\ell^2\omega^2$$

### Basic setup

Here are the basic parameters we are going to use for this exercise:

```python
g = 9.81   # m/s^2
l = 0.5    # length of pendulum, in meters
tmax = 50. # seconds
t = np.linspace(0, tmax, int(100*tmax))
```

Write a function `derivs` for usage with `scipy.integrate.odeint` that computes the derivatives for the damped, driven harmonic oscillator. The solution vector at each time will be $\vec{y}(t) = (\theta(t),\omega(t))$.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def derivs(y, t, a, b, omega0):
    """Compute the derivatives of the damped, driven pendulum.
    
    Parameters
    ----------
    y : ndarray
        The solution vector at the current time t[i]: [theta[i],omega[i]].
    t : float
        The current time t[i].
    a, b, omega0: float
        The parameters in the differential equation.
    
    Returns
    -------
    dy : ndarray
        The vector of derviatives at t[i]: [dtheta[i],domega[i]].
    """
    if y.shape == (2,):
        domega = (-9.81/.5)*np.sin(y[0])-a*y[1]-b*np.sin(omega0*t)
        dtheta = y[1]
    else:
        domega = (-9.81/.5)*np.sin(y[:,0])-a*y[:,1]-b*np.sin(omega0*t)
        dtheta = y[:,1]
    dy = np.column_stack((dtheta,domega))
    return dy
```

```python
a = np.array([[0,1],[2,0]])
derivs(np.array([np.pi,1.0]), 0, 1.0, 1.0, 1.0)
.5*.5**2
```

Output:
```
0.125
```

<!-- grade_id: odesex03a, points: 2, grade: true, solution: false -->
```python
assert np.allclose(derivs(np.array([np.pi,1.0]), 0, 1.0, 1.0, 1.0), [1.,-1.])
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def energy(y):
    """Compute the energy for the state array y.
    
    The state array y can have two forms:
    
    1. It could be an ndim=1 array of np.array([theta,omega]) at a single time.
    2. It could be an ndim=2 array where each row is the [theta,omega] at single
       time.
    
    Parameters
    ----------
    y : ndarray, list, tuple
        A solution vector
    
    Returns
    -------
    E/m : float (ndim=1) or ndarray (ndim=2)
        The energy per mass.
    """
    if y.shape == (2,):
        E = 9.81*.5*(1-np.cos(y[0])) + (.5**3)*(y[1]**2)
    else:
        E = 9.81*.5*(1-np.cos(y[:,0])) + (.5**3)*(y[:,1]**2)
    return E
        
```

<!-- grade_id: odesex03b, points: 2, grade: true, solution: false -->
```python
assert np.allclose(energy(np.array([np.pi,0])),g)
assert np.allclose(energy(np.ones((10,2))), np.ones(10)*energy(np.array([1,1])))
```

### Simple pendulum

Use the above functions to integrate the simple pendulum for the case where it starts at rest pointing vertically upwards. In this case, it should remain at rest with constant energy.

* Integrate the equations of motion.
* Plot $E/m$ versus time.
* Plot $\theta(t)$ and $\omega(t)$ versus time.
* Tune the `atol` and `rtol` arguments of `odeint` until $E/m$, $\theta(t)$ and $\omega(t)$ are constant.

Anytime you have a differential equation with a a conserved quantity, it is critical to make sure the numerical solutions conserve that quantity as well. This also gives you an opportunity to find other bugs in your code. The default error tolerances (`atol` and `rtol`) used by `odeint` are not sufficiently small for this problem. Start by trying `atol=1e-3`, `rtol=1e-2` and then decrease each by an order of magnitude until your solutions are stable.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
y0 = np.array([0,0])
a, b =0.0, 0.0
omega0=0.0
a = odeint(derivs,y0,t, args=(a,b,omega0))

```

Output:
```
<ipython-input-134-58e7270489c3> in <module>()
      2 a, b =0.0, 0.0
      3 omega0=0.0
----> 4 a = odeint(derivs,y0,t, args=(a,b,omega0))

/usr/local/lib/python3.4/dist-packages/scipy/integrate/odepack.py in odeint(func, y0, t, args, Dfun, col_deriv, full_output, ml, mu, rtol, atol, tcrit, h0, hmax, hmin, ixpr, mxstep, mxhnil, mxordn, mxords, printmessg)
    146     output = _odepack.odeint(func, y0, t, args, Dfun, col_deriv, ml, mu,
    147                              full_output, rtol, atol, tcrit, h0, hmax, hmin,
--> 148                              ixpr, mxstep, mxhnil, mxordn, mxords)
    149     if output[-1] < 0:
    150         print(_msgs[output[-1]])

ValueError: object too deep for desired array
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: odesex03c, points: 3, grade: true, solution: false -->
```python
assert True # leave this to grade the two plots and their tuning of atol, rtol.
```

## Damped pendulum

Write a `plot_pendulum` function that integrates the damped, driven pendulum differential equation for a particular set of parameters $[a,b,\omega_0]$.

* Use the initial conditions $\theta(0)=-\pi + 0.1$ and $\omega=0$.
* Decrease your `atol` and `rtol` even futher and make sure your solutions have converged.
* Make a parametric plot of $[\theta(t),\omega(t)]$ versus time.
* Use the plot limits $\theta \in [-2 \pi,2 \pi]$ and $\theta \in [-10,10]$
* Label your axes and customize your plot to make it beautiful and effective.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def plot_pendulum(a=0.0, b=0.0, omega0=0.0):
    """Integrate the damped, driven pendulum and make a phase plot of the solution."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

Here is an example of the output of your `plot_pendulum` function that should show a decaying spiral.

```python
plot_pendulum(0.5, 0.0, 0.0)
```

Use `interact` to explore the `plot_pendulum` function with:

* `a`: a float slider over the interval $[0.0,1.0]$ with steps of $0.1$.
* `b`: a float slider over the interval $[0.0,10.0]$ with steps of $0.1$.
* `omega0`: a float slider over the interval $[0.0,10.0]$ with steps of $0.1$.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Use your interactive plot to explore the behavior of the damped, driven pendulum by varying the values of $a$, $b$ and $\omega_0$.

* First start by increasing $a$ with $b=0$ and $\omega_0=0$.
* Then fix $a$ at a non-zero value and start to increase $b$ and $\omega_0$.

Describe the different *classes* of behaviors you observe below.

<!-- grade_id: odesex03d, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE